```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import geopandas as gpd

import ipywidgets as widgets
from ipywidgets import interact, interactive, fixed, interact_manual
import plotly.express as px
import plotly.graph_objects as go
```

```python
from urllib.request import urlopen
import json
with urlopen('https://raw.githubusercontent.com/plotly/datasets/master/geojson-counties-fips.json') as response:
    counties = json.load(response)
```

```python
counties
```

```python
# load the data
df = pd.read_csv("data/rgdp_county_allindustry.csv")
# converting GeoFIPS to string and fill with leading zeros to match the GitHUB data
df['GeoFIPS'] = df['GeoFIPS'].astype(str).str.zfill(5) 
df.head()
```

```python
# calculate the percentage change
df_pct_chg = df.copy()
df_pct_chg = df_pct_chg.sort_values(by=['GeoFIPS', 'year']).reset_index(drop=True)
df_pct_chg['pct_chg'] = df_pct_chg['value'].pct_change()*100
df_pct_chg = df_pct_chg[df_pct_chg[year] != 2001].reset_index(drop=True)
df_pct_chg.head()
```

```python
df_pct_chg2019 = df_pct_chg[df_pct_chg['year'] == 2019].reset_index(drop=True)
```

```python
fig = px.choropleth_mapbox(df_pct_chg2019, geojson=counties, locations='GeoFIPS', color='pct_chg',
                           color_continuous_scale="RdBu_r",
                           range_color=(0, 12),
                           mapbox_style="carto-positron",
                           zoom=3, center = {"lat": 37.0902, "lon": -95.7129},
                           opacity=0.5,
                           labels={'pct_chg':'change in GDP (%)'}
                          )
fig.update_layout(margin={"r":0,"t":0,"l":0,"b":0})
fig.show()
```

```python
# Create a list of dictionaries, each representing a different year's data
data = []
for year in df_pct_chg['year'].unique():
    data.append(
        dict(
            args=['visible', [False]*len(df_pct_chg['year'].unique())],
            label=str(year),
            method='update'
        )
    )

# Update the visibility for each year
for i in range(len(data)):
    year = data[i]['label']
    data[i]['args'][1][i] = True  # Toggle i'th trace to "visible"

# Create the figure with a dropdown menu
fig = go.Figure()

for year in df_pct_chg['year'].unique():
    fig.add_trace(
        go.Choroplethmapbox(
            geojson=counties,
            locations=df_pct_chg[df_pct_chg['year'] == year]['GeoFIPS'],
            z=df_pct_chg[df_pct_chg['year'] == year]['pct_chg'],
            colorscale="Viridis",
            zmin=0,
            zmax=12,
            marker_opacity=0.5,
            marker_line_width=0,
        )
    )

fig.update_layout(
    updatemenus=[
        dict(
            buttons=list(data),
            direction="down",
            pad={"r": 10, "t": 10},
            showactive=True,
            x=0.1,
            xanchor="left",
            y=1.1,
            yanchor="top"
        ),
    ]
)

fig.update_layout(mapbox_style="carto-positron", mapbox_zoom=3, mapbox_center = {"lat": 37.0902, "lon": -95.7129})
fig.update_layout(margin={"r":0,"t":0,"l":0,"b":0})
fig.show()
```

```python
# Create a list of dictionaries, each representing a different year's data
MIN = np.nanpercentile(df_pct_chg["pct_chg"], 10)
MAX = np.nanpercentile(df_pct_chg["pct_chg"], 90)
data = []
for i, year in enumerate(df_pct_chg['year'].unique()):
    visible = [False]*len(df_pct_chg['year'].unique())
    visible[i] = True
    data.append(
        dict(
            args=[{'visible': visible}],
            label=str(year),
            method='update'
        )
    )

# Create the figure with a dropdown menu
fig = go.Figure()

for year in df_pct_chg['year'].unique():
    fig.add_trace(
        go.Choroplethmapbox(
            geojson=counties,
            locations=df_pct_chg[df_pct_chg['year'] == year]['GeoFIPS'],
            z=df_pct_chg[df_pct_chg['year'] == year]['pct_chg'],
            colorscale="Viridis",
            zmin=MIN,
            zmax=MAX,
            marker_opacity=0.5,
            marker_line_width=0,
            visible = True if i == 0 else False
        )
    )

fig.update_layout(
    updatemenus=[
        dict(
            buttons=list(data),
            direction="down",
            pad={"r": 10, "t": 10},
            showactive=True,
            x=0.1,
            xanchor="right",
            y=1.1,
            yanchor="top"
        ),
    ],
    autosize=False,
    width=1000,
    height=800
)

fig.update_layout(mapbox_style="carto-positron", mapbox_zoom=3, mapbox_center = {"lat": 37.0902, "lon": -95.7129})
fig.update_layout(margin={"r":0,"t":0,"l":0,"b":0})
fig.show()
```

Tests for q10 (points: None)
```
>>> assert sum([ord(c) for c in codeword]) ^ 966984 == 966771
```


Tests for q1_1 (points: None)
```
>>> assert rgdp_raw.shape == (47674, 29)
```


Tests for q1_2 (points: None)
```
>>> assert rgdp.shape == (47670, 29)
```
```
>>> assert rgdp.iloc[-1]['GeoName'] != np.nan
```
```
>>> assert '98000' in rgdp.iloc[-1]['GeoFIPS']
```


Tests for q2_1 (points: None)
```
>>> assert isinstance(N_rows, int) and N_rows > 0
```
```
>>> assert isinstance(N_cols, int) and N_cols > 0
```


Tests for q2_2 (points: None)
```
>>> assert 100 < N_unique_geofips < 10000
```


Tests for q3_2 (points: None)
```
>>> assert rgdp_no_nans.shape == (34206, 29)
```
```
>>> assert not rgdp_no_nans.isnull().values.any()
```
failure_message: The dataframe should not contain any NaNs.
```
>>> assert not '(D)' in rgdp_no_nans.values
```
failure_message: The dataframe should not contain any "(D)"s.
```
>>> assert not '(NA)' in rgdp_no_nans.values
```
failure_message: The dataframe should not contain any "(NA)"s.


Tests for q4_1 (points: None)
```
>>> assert np.all([rgdp_clean[str(year)].dtype == 'int64' for year in np.arange(2001, 2022)])
```


Tests for q4_2 (points: None)
```
>>> assert no_quotation_marks == 'UC Berkeley is the No.1 university in California.'
```


Tests for q4_3 (points: None)
```
>>> assert rgdp_clean['GeoFIPS'].dtypes == 'int64'
```
```
>>> assert np.isclose(len(rgdp_clean['GeoFIPS'].unique()), 3166)
```


Tests for q5_1 (points: 3)
```
>>> assert rgdp_melted.shape == (718326, 7)
```
```
>>> assert np.all(rgdp_melted.columns == ['GeoFIPS', 'GeoName', 'Region', 'LineCode', 'Description', 'year', 'value'])
```
```
>>> assert np.isclose(rgdp_melted[(rgdp_melted['year'] == '2001') & (rgdp_melted['GeoFIPS'] == 0) & (rgdp_melted['LineCode'] == 1)]['value'].iloc[0], 13263417000)
```


Tests for q5_2 (points: None)
```
>>> assert rgdp_melted['year'].dtype == 'int64'
```


Tests for q5_3 (points: None)
```
>>> assert rgdp_county.shape == (699426, 7)
```
```
>>> assert np.all(rgdp_county.columns == ['GeoFIPS', 'GeoName', 'Region', 'LineCode', 'Description', 'year', 'value'])
```
```
>>> assert np.isclose(rgdp_county[(rgdp_county['year'] == 2020) & (rgdp_county['GeoFIPS'] == 56045)]['value'].iloc[0], 292568)
```


Tests for q5_4 (points: None)
```
>>> assert rgdp_county_allindustry.shape == (65226, 6)
```
```
>>> assert np.all(rgdp_county_allindustry.columns == ['GeoFIPS', 'GeoName', 'Region', 'LineCode', 'year', 'value'])
```
```
>>> assert np.min(rgdp_county_allindustry['LineCode']) == np.max(rgdp_county_allindustry['LineCode'])
```
```
>>> assert np.isclose(rgdp_county_allindustry[(rgdp_county_allindustry['year'] == 2020) & (rgdp_county_allindustry['GeoFIPS'] == 56045)]['value'].iloc[0], 292568)
```


Tests for q6_1 (points: None)
```
>>> assert rgdp_county_allindustry_mean.shape == (21, 2)
```
```
>>> assert np.all(rgdp_county_allindustry_mean.columns == ['year', 'value'])
```


Tests for q6_2 (points: None)
```
>>> assert rgdp_county_allindustry_median.shape == (21, 2)
```
```
>>> assert np.all(rgdp_county_allindustry_median.columns == ['year', 'value'])
```


Tests for q7_1 (points: None)
```
>>> assert rgdp_far_west_pct_chg.shape == (20, 2)
```
```
>>> assert np.all(~rgdp_far_west_pct_chg['value'].isnull())
```
```
>>> assert rgdp_far_west_pct_chg['value'].iloc[0] < 1
```


Tests for q8_1 (points: None)
```
>>> assert county_geo.shape == (3143, 4)
```
```
>>> assert np.all(county_geo.columns == ['county_fips', 'lat', 'lng', 'population'])
```
```
>>> assert county_geo[county_geo['county_fips'] == 17031]['lat'].iloc[0] == 41.8401
```


Tests for q8_2 (points: None)
```
>>> assert rgdp_county_allindustry_geo.shape == (64722, 10)
```
```
>>> assert np.all([c in rgdp_county_allindustry_geo.columns for c in ['GeoFIPS', 'GeoName', 'Region', 'LineCode', 'year', 'value', 'county_fips', 'lat', 'lng', 'population']])
```
