```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

```python
import pandas as pd
import numpy as np
import statsmodels.api as sm
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline
```

# Lab 7: The Impact of Minimum Wage on Employment

## Overview

We are going to replicate a study conducted by **[Card and Krueger in 1994](https://davidcard.berkeley.edu/papers/njmin-aer.pdf)** that investigates the relationship between a rise in the minimum wage and employment.

## Background

[Economic theories](https://www.frbsf.org/research-and-insights/publications/economic-letter/2015/12/effects-of-minimum-wage-on-employment/) have long suggested that increases in the minimum wage lead to a reduction in employment for at least two reasons:

1. **Businesses are less likely to hire** and will rather invest in other resources that are now cheaper because of the wage increase.
2. **Higher salaries will induce businesses to raise their prices** to compensate for their greater costs; as prices increase, we expect fewer buyers, which will lead to lower demand and employment.

These theories have found [mixed support](https://www.nber.org/papers/w12663), but the discussion is still very much open within the policy world, as states discuss the opportunity to raise their minimum wage to help local populations to face increasing living costs. Discussions are currently occurring in **[New Jersey](https://www.nytimes.com/2019/01/17/nyregion/nj-minimum-wage.html)** and **[Illinois](https://kmox.radio.com/articles/discussions-underway-raise-illinois-minimum-wage-15hour)** to raise the minimum wage to **$15/hour** ([New York](https://www.nytimes.com/2018/12/31/nyregion/15-minimum-wage-new-york.html?module=inline) has successfully passed this same raise in 2018).

## The Original Study

One of the first studies looking at this policy problem was **Card and Krueger’s**. They applied a difference-in-difference design to look at two groups of fast-food restaurants:

- Fast-food restaurants in **New Jersey** where the minimum wage increased from \\$4.25 to $5.05 per hour (treatment group)
- Fast-food restaurants in **Pennsylvania** where the minimum wage did not increase (control group).

They collected data before and after the minimum wage was approved. Data used in the study can be downloaded [here](https://github.com/DS4PS/PROG-EVAL-III/blob/master/TEXTBOOK/DATA/DID_Example.csv).

## Research Question

- **Do increases in the minimum wage affect employment?**

## Hypothesis

- **An increase in the minimum wage is negatively correlated with employment.**


## Part 1: Diff-in-Diff Conceptual

<!-- BEGIN QUESTION -->

**Question 1.1:** How does the choice of research design, specifically a difference-in-differences approach, influence the validity and reliability of conclusions drawn from empirical studies examining the impact of policy changes, such as minimum wage increases, on economic outcomes like employment? Hint: Think about the parallell trend assumption.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 2: Working with the Data

We now proceed with loading in the data used in the original 1994 study. It cointains the following variables:

| Variable Name | Description                                          |
|---------------|------------------------------------------------------|
| ID            | Unique identifier for fast food                     |
| Treatment     | Pre-treatment (=0) and post-treatment (=1)          |
| Group         | 1 if NJ (treatment); 0 if PA (Control)              |
| Empl          | # of full time employees                             |
| C.Owned       | If owned by a company (=1) or not (=0)               |
| Hours.Opening | Number hours open per day                            |
| Soda          | Price of medium soda, including tax                  |
| Fries         | Price of small fries, including tax                  |
| Chain         | 1 = BK, 2 = KFC, 3 = Roys, 4 = Wendys                |
| SouthJ        | South New Jersey                                     |
| CentralJ      | Central New Jersey                                   |
| NorthJ        | North New Jersey                                     |
| PA1           | Northeast suburbs of Philadelphia                    |
| PA2           | Easton and other PA areas                            |
| Shore         | New Jersey Shore                                     |


```python
dd_df = pd.read_csv('data/DID_Example.csv')
dd_df
```

**Question 2.1:** Summarize the data above using the `.describe()` Pandas function.

```python

dd_df_summarized = ...


dd_df_summarized
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Question 2.2:** What does each row in `dd_df` represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.3:** What's the average number of employees per fast-food restaurant?

```python
mean_employees = ...

mean_employees
```

```python
grader.check("q2_3")
```

**Question 2.4:** What percentage of fast-food restaurants are part of the treatment group?

```python
per_treatment = ...

per_treatment
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Question 2.5:** What's the difference between the `Group` the `Treatment` columns?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.6:** Notice that our `Chain` column is a categorical variable hidden as a numeric value! One-hot encode this column and save your result in `dd_df_encoded`. You should drop the original column and create 3 new columns: 'BK', KFC', and 'Wendys'. This is helpful for our analysis below.

```python
chain_dummies = ...

...

dd_df_encoded = ...

dd_df_encoded
```

```python
grader.check("q2_6")
```

## Part 3: Analysis: Estimating the Difference-in-Differences Model

We now estimate the difference-in-differences model based on the model below. Along with the **Group** and **Treatment** variables, we also include a set of control variables to account for differences across restaurants. For instance, we consider the variable **opening hours**—suggesting that fast-food restaurants open for more hours might need more employees. Similarly, we account for the prices of **fries and sodas**, under the assumption that more expensive fast food might have more resources to hire additional staff.


<!-- BEGIN QUESTION -->

**Question 3.1:** Using the general diff-in-diff model below, specify the model we're about to estimate using $\LaTeX$. Hint: Don't forget to add your controls.

$(1.1)$ 

$$Y = \beta_0 + \beta_1 \cdot \text{Treatment} + \beta_2 \cdot \text{Post} + \beta_3 \cdot \text{Treatment} \times \text{Post} + \text{Controls} + e$$


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.2:** Now, rename the relevant column names of the `dd_df` to fit the equation specified in (1.1). Your new column names should be `Treatment` and `Post`, both of which take on binary values. This will simplify our coming analysis!

```python
dd_df = dd_df_encoded...
dd_df
```

```python
grader.check("q3_2")
```

**Question 3.3:** Now, using `statsmodels` run the diff-in-diff from equation (1.1). Note: This is without your controls.

```python

# Create your Treatment*Post variable.
dd_df['Treatment_Post'] = ...


# Define your dependent variable.
Y = dd_df[...] 


# Define your independent variables, including the interaction term
X = dd_df[...]


# Add a constant to the model (the intercept)
X = sm.add_constant(X)

# Create the OLS model
model = ...


# Fit the model
results = ...

# Print the summary of the model
print(results.summary())


```

```python
grader.check("q3_3")
```

**Question 3.4.1:** You notice that your output from the `print(results.summary())` gives you all NaN's as cofficients and std errors. This may be because of NaNs in the data. For the sake of simplicity, feel free to just drop all the NaNs for now.

```python
dd_df = ...
dd_df
```

```python
grader.check("q3_4_1")
```

<!-- BEGIN QUESTION -->

**Question 3.4.2:** Describe in 1-2 sentences below what you did and how it might affect your later analysis. Also note how many observations you lost. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.5:** Repeat the analysis from 3.3 and estimate the diff-in-diff estimator.

```python
...
```

```python
grader.check("q3_5")
```

<!-- BEGIN QUESTION -->

**Question 3.6:** Using the Sign, Size, and Significance framework as described below, interpret your findings. Please format your markdown nicely (like the following cell) to aid readability.

### Sign, Size, and Significance (SSS) framework for interpreting regression outputs

#### 1. Sign

- **Expected Sign**: What sign did you expect the estimated parameter(s) to have? Why?
- **Actual Sign**: Does your estimate(s) have this sign (i.e., are you surprised or reassured by your results)?

#### 2. Significance

- **Statistical Significance**: Is the estimate(s) statistically different from zero?
- **T-Statistic**: What is the t-statistic of this hypothesis?

#### 3. Size

- **Effect on Dependent Variable**: How do changes in this variable affect the dependent variable according to your estimation?
- **Economic Significance**: Is this an economically meaningful effect size?

This framework is borrowed from Berkeley's EEP C118 course. See more [here.](https://are.berkeley.edu/courses/EEP118/spring2014/section/Handout4_2014.pdf)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.7:** We notice that our $R^2$ (one of many goodness-of-fit parameters we can use) is very low. To aid our analysis, Include the controls you specified in the $\LaTeX$ equation above.

```python
...
```

```python
grader.check("q3_7")
```

<!-- BEGIN QUESTION -->

**Question 3.8:** Controlling for opening hours, and prices of soda & fries - What is the diff-in-diff estimator for the impact on introducing a minimum wage in New Jersey? Remember to include units and give a brief interpretation of our findings following the SSS framework from above. Be sure to mention how your $R^2$ changed, and it's practical implications. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 4: Visualizing the diff-in-diff estimator

Below, we plot the difference in difference estimates from question 3.5 and embark on a visual interpretation of our findings.

```python
# First, calculate mean employment by group and time
mean_empl = dd_df.groupby(['Treatment', 'Post'])['Empl'].mean().reset_index()

# Now, plot these averages with lines to show the change from pre to post for each group
plt.figure(figsize=(14, 7))
sns.lineplot(x='Post', y='Empl', hue='Treatment', data=mean_empl, marker='o', palette='viridis')

plt.title('Difference-in-Differences: Employment Pre and Post Treatment')
plt.xlabel('Period (0 = Pre-Treatment, 1 = Post-Treatment)')
plt.ylabel('Average Employment')
plt.xticks(ticks=[0, 1], labels=['Pre-Treatment', 'Post-Treatment'])
plt.legend(title='Group / State', labels=['Control / PA', 'Treatment / NJ'])

# Adding annotations for clarity
for line in range(mean_empl.shape[0]):
    plt.text(mean_empl.Post[line]+0.02, mean_empl.Empl[line], 
             f"{mean_empl.Empl[line]:.2f}", horizontalalignment='left', 
             size='medium', color='black', weight='semibold')

plt.tight_layout()
plt.show();
```

**Question 4.1:** Now, let's make sure we can visually interpret our findings. Match the regression output numbers below to the appropriate regression equation.

**Equations:**

a) $const$

b) $const + \text{Treatment}$

c) $const + \text{Post}$

d) $const + \text{Post} + \text{Treatment} + \text{Treatment*Post}$

**Regression Output Numbers:**

1. $10.44$
2. $7.59$
3. $7.75$
4. $8.49$


```python

a = ... # These should be either 10.44, or 7.59, or 7.75, or 8.49
b = ...
c = ...
d = ...



```

```python
grader.check("q4_1")
```

**Question 4.2:** Using the numbers from the visualization above, calculate the difference in difference estimate. It should match your estimate from question 3.5.

```python
DD_estimate = ...
DD_estimate
```

```python
grader.check("q4_2")
```

<!-- BEGIN QUESTION -->

**Question 4.3:** Looking back to your response to question 1.1 and the plot from above, what is a central piece of evidence we're lacking from dataset that would strengthen the validity of our diff-in-diff analysis?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Conclusion
Congratulations, you have finished lab 7! We hope you enjoyed the lab - you're one step closer in becoming a master replicator of economics papers...

Have a great week(end)!

Peter.


---
## Feedback

**Question 5:** Please fill out this short [feedback form](https://forms.gle/58k2xRHNn6FSyAJ1A) to let us know your thoughts about this lab! We really appreciate your opinions and feedback! At the end of the Google form, you should see a codeword. Assign the codeword to the variable `codeword` below. 

```python
codeword = ...
```

```python
grader.check("q5")
```

---
## Sources

*FOUNDATIONS OF PROGRAM EVALUATION III REGRESSION TOOLS FOR CAUSAL ANALYSIS* by Data Science for 
Public Service (https://ds4ps.org/PROG-EVAL-III/index.html), retrieved 22 Feb 2024.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2_1 (points: 1)
```
>>> assert dd_df_summarized.shape == (8, 15)
```


Tests for q2_3 (points: 1)
```
>>> assert np.isclose(mean_employees, 8.238778054862843)
```


Tests for q2_4 (points: 1)
```
>>> assert np.isclose(per_treatment, 80.73170731707317)
```


Tests for q2_6 (points: 2)
```
>>> assert dd_df_encoded.shape == (820, 17)
```
```
>>> expected_columns = ['BK', 'KFC', 'Wendys']
>>> assert all((column in dd_df_encoded.columns for column in expected_columns))
```


Tests for q3_2 (points: 1)
```
>>> expected_columns_after_rename = ['Treatment', 'Post']
>>> assert all((column in dd_df.columns for column in expected_columns_after_rename))
```
```
>>> original_columns = ['Group']
>>> assert all((column not in dd_df.columns for column in original_columns))
```


Tests for q3_3 (points: 2)
```
>>> assert np.isnan(results_33.params.const)
```
```
>>> assert len(results_33.params) == 4
```


Tests for q3_4_1 (points: 1)
```
>>> expected_shape = (757, 18)
>>> assert dd_df.shape == expected_shape
```
```
>>> assert not dd_df.isna().any().any()
```


Tests for q3_5 (points: 2)
```
>>> expected_coefficients = {'const': {'coef': 10.4408, 'se': 0.929, 'ci_lower': 8.616, 'ci_upper': 12.265}, 'Treatment': {'coef': -2.8472, 'se': 1.037, 'ci_lower': -4.883, 'ci_upper': -0.811}, 'Post': {'coef': -2.6943, 'se': 1.337, 'ci_lower': -5.32, 'ci_upper': -0.069}, 'Treatment_Post': {'coef': 3.5708, 'se': 1.49, 'ci_lower': 0.646, 'ci_upper': 6.495}}
>>> tolerance = 0.001
>>> for var, expected in expected_coefficients.items():
...     coef, se = (results_35.params[var], results_35.bse[var])
...     ci_lower, ci_upper = results_35.conf_int().loc[var]
...     assert abs(coef - expected['coef']) < tolerance, f"Coefficient for {var} is incorrect: expected {expected['coef']}, got {coef}"
...     assert abs(se - expected['se']) < tolerance, f"Standard error for {var} is incorrect: expected {expected['se']}, got {se}"
...     assert abs(ci_lower - expected['ci_lower']) < tolerance, f"Lower CI for {var} is incorrect: expected {expected['ci_lower']}, got {ci_lower}"
...     assert abs(ci_upper - expected['ci_upper']) < tolerance, f"Upper CI for {var} is incorrect: expected {expected['ci_upper']}, got {ci_upper}"
```


Tests for q3_7 (points: 2)
```
>>> expected_coefficients = {'const': {'coef': 1.247, 'se': 3.451, 'ci_lower': -5.528, 'ci_upper': 8.022}, 'Treatment': {'coef': -2.9825, 'se': 1.056, 'ci_lower': -5.055, 'ci_upper': -0.91}, 'Post': {'coef': -1.4365, 'se': 1.3, 'ci_lower': -3.989, 'ci_upper': 1.116}, 'Treatment_Post': {'coef': 4.1275, 'se': 1.435, 'ci_lower': 1.311, 'ci_upper': 6.944}, 'Hours.Opening': {'coef': 0.5209, 'se': 0.071, 'ci_lower': 0.382, 'ci_upper': 0.659}, 'Fries': {'coef': -0.3659, 'se': 3.549, 'ci_lower': -7.333, 'ci_upper': 6.602}, 'Soda': {'coef': 2.0506, 'se': 4.222, 'ci_lower': -6.238, 'ci_upper': 10.339}}
>>> tolerance = 0.001
>>> for var, expected in expected_coefficients.items():
...     coef, se = (results_37.params[var], results_37.bse[var])
...     ci_lower, ci_upper = results_37.conf_int().loc[var]
...     assert abs(coef - expected['coef']) < tolerance, f"Coefficient for {var} is incorrect: expected {expected['coef']}, got {coef}"
...     assert abs(se - expected['se']) < tolerance, f"Standard error for {var} is incorrect: expected {expected['se']}, got {se}"
...     assert abs(ci_lower - expected['ci_lower']) < tolerance, f"Lower CI for {var} is incorrect: expected {expected['ci_lower']}, got {ci_lower}"
...     assert abs(ci_upper - expected['ci_upper']) < tolerance, f"Upper CI for {var} is incorrect: expected {expected['ci_upper']}, got {ci_upper}"
```


Tests for q4_1 (points: 1)
```
>>> assert np.isclose(np.mean([a, b]), 9.015)
>>> assert np.isclose(np.mean([b, c]), 7.67)
>>> assert np.isclose(np.mean([c, d]), 8.120000000000001)
```


Tests for q4_2 (points: 1)
```
>>> assert np.isclose(DD_estimate, 3.5708, 0.001)
```


Tests for q5 (points: None)
```
>>> assert sum([ord(c) for c in codeword]) ^ 966984 == 966706
```
